# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # YOUR CODE HERE
    #raise NotImplementedError()
    c=0
    x=str(x)
    y=str(y)
    for i in range(len(x)):
        if len(x)!=len(y):
            break
        if x[i]==y[]:
            c+=1
    if c == len(x):
        print('True')
        return True
```

Output:
```
True
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
import numpy as np
n =[]
def find(z):
    y=range(1,7)
    num=[z*i for i in y]
    print (num)
    for i in range(len(num)-1):
        if same_digits(str(num[i]),str(num[i+1])) is True:
            n.append(z)
    return(min(n))
find(22)
```

Output:
```
[22, 44, 66, 88, 110, 132]
<ipython-input-1-6740b5d3ad46> in <module>()
     11             n.append(z)
     12     return(min(n))
---> 13 find(22)

<ipython-input-1-6740b5d3ad46> in find(z)
      8     print (num)
      9     for i in range(len(num)-1):
---> 10         if same_digits(str(num[i]),str(num[i+1])) is True:
     11             n.append(z)
     12     return(min(n))

NameError: name 'same_digits' is not defined
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```